Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your github nickname below:

```python
NICKNAME = "tttonyalpha"
```

---

Write a function which solves a quadratic equation, $x^2 + b x + c = 0$. Your function must receive two values for coefficients and return two real or complex roots. Always return two roots, even if the equation has a single double root. 

<!-- grade_id: cell-2032efc64a327acc, points: 0, grade: false, solution: true -->
```python
from cmath import sqrt

def solve_quad(b, c):
    ''' Parameters
    ----------
    b, c : float
       Coefficients
       
    Returns
    -------
    x1, x2 : float or complex
        Roots.
    '''
    d = sqrt(b**2 - 4*c)
    
    x1 = (- b + d) / 2
    x2 = (- b - d) / 2
    
    #b >> c: 
    if (c != 0) and (x1 == 0 or x2 == 0):
        x1 = - c / b
        x2 = - b
        
    return x1, x2
```

<!-- grade_id: cell-89cf7badb387b0e0, points: 5, grade: true, solution: false -->
```python
#
# Your implementation should pass tests in this cell.
#
# Do not remove or alter this cell. You may run it, but do not remove or alter it.
# Your changes to this cell will be ignored on grading.
# There will be additional tests.
#
from numpy.testing import assert_allclose

variants = [{'b': 4.0, 'c': 3.0},
            {'b': 2.0, 'c': 1.0},
            {'b': 0.5, 'c': 4.0},
            {'b': 1e10, 'c': 3.0},
           ]

for var in variants:
    x1, x2 = solve_quad(**var)
    assert_allclose(x1*x2, var['c'], rtol=1e-15)
```

<!-- grade_id: cell-0970e4298cb77486, points: 5, grade: true, solution: false -->
```python

```